# MSCF 46982 Market Microstructure and Algorithmic Trading
# Fall 2018 Mini 2

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", specify your name and that of your collaborator, and remove the `notimplemented` exception.



---

```q
NAME: "Ze Yang"
COLLABORATOR: ""
```

## Anaconda w/ Embedpy Installation

To confirm you have installed Anaconda with the Kdb+ kernel installed properly, compute the `md5` hash of 3.6.

### Part A (1 points)

The `.z.K` internal variable contains the version of Kdb+. Evaluate the below cell, to confirm you are running version 3.6 and display the md5 hashed value of this version. The code below has the hashed value of 3.5. Update the declaration of `x` to have the correct hashed value of 3.6.

<!-- grade_id: dotzdotk, points: 0, grade: false, solution: true -->
```q
.z.K
md5 string .z.K
x:0x804066a6a5042090a2a0b25c55f38e72
/ YOUR CODE HERE
x:0xe44436592c2daabbd3d797f8967245a3
x
```

Output:
```
3.6

0xe44436592c2daabbd3d797f8967245a3

0xe44436592c2daabbd3d797f8967245a3
```

<!-- grade_id: correct_dotzdotk, points: 1, grade: true, solution: false -->
```q
assert:{if[not x~y;'`$"expecting '",(-3!x),"' but found '",(-3!y),"'"]}
assert[md5 "3.6"] x
```

## Fibonacci Sequence
While Kdb+ comes with a `while` operator, it is rarely used and frowned upon as good form.  This problem takes you through a popular example of how iteration is performed in q using adverbs. A quick introduction to Kdb+ and an example of the Fibonacci sequence can also be view in [Q for All](https://code.kx.com/q/tutorials/q-for-all).

### Part A (1 points)

Write a function that returns the sum of the last 2 elements of a list.

Hint: use the `#` operator.

<!-- grade_id: sumlasttwo, points: 0, grade: false, solution: true -->
```q
/ returns the sum of the last two elements
sumlasttwo:{
    / YOUR CODE HERE
    x:sum -2#x;
    /'`notimplemented
    x}

```

Your function should return `8` for `x:0 1 1 2 3 5`. Check that it does:

```q
sumlasttwo 0 1 1 2 3 5
```

Output:
```
8
```

<!-- grade_id: correct_sumlasttwo, points: 1, grade: true, solution: false -->
```q
/ check that sumlasttwo produces the correct output for several inputs
assert[2] sumlasttwo 0 1 1
assert[3] sumlasttwo 0 1 1 2
assert[5] sumlasttwo 0 1 1 2 3
assert[8] sumlasttwo 0 1 1 2 3 5
assert[13] sumlasttwo 0 1 1 2 3 5 8
```

### Part B (1 point)

Write a function `fib` that appends the sum of the last two elements of a list to the list.

<!-- grade_id: fib, points: 0, grade: false, solution: true -->
```q
/ returns the sum of the last two elements
fib:{
    / YOUR CODE HERE
    x:x,sumlasttwo x;
    x}

```

Your function should return `0 1 1 2 3 5 8` for `x:0 1 1 2 3 5`.  Check that it does:

```q
fib 0 1 1 2 3 5
```

Output:
```
0 1 1 2 3 5 8
```

<!-- grade_id: correct_fib, points: 1, grade: true, solution: false -->
```q
/ check that fib produces the correct output for several inputs
assert[0 1 1] fib 0 1
assert[0 1 1 2] fib 0 1 1
assert[0 1 1 2 3] fib 0 1 1 2
assert[0 1 1 2 3 5] fib 0 1 1 2 3
assert[0 1 1 2 3 5 8] fib 0 1 1 2 3 5
assert[0 1 1 2 3 5 8 13] fib 0 1 1 2 3 5 8
```

### Part C (1 point)

Using your `fib` function and the `/` [adverb](https://code.kx.com/q/ref/adverbs/#converge-repeat), write a new function `fibonacci` which returns the first n elements of the Fibonacci sequence after `0 1`.

<!-- grade_id: fibonacci, points: 0, grade: false, solution: true -->
```q
/ returns the sum of the last two elements
fibonacci:{[n]
    / YOUR CODE HERE
    x:fib/[n;0 1];
    x}

```

Your function should return `0 1 1 2 3 5` for `n:4`.  Check that it does:

```q
fibonacci 4
```

Output:
```
0 1 1 2 3 5
```

<!-- grade_id: correct_fibonacci, points: 1, grade: true, solution: false -->
```q
/ check that fibonacci produces the correct output for several inputs
assert[0 1] fibonacci 0
assert[0 1 1] fibonacci 1
assert[0 1 1 2] fibonacci 2
assert[0 1 1 2 3] fibonacci 3
assert[0 1 1 2 3 5] fibonacci 4
assert[0 1 1 2 3 5 8] fibonacci 5
assert[1298777728820984005] last fibonacci 100
```

## Sharpe Ratio

A common performance statistic for investors is the **sharpe ratio**.  It is defined as the average profit (above the risk free rate) divided by the risk (or standard deviation) of achieving that profit.

$$\frac{\text{average daily return} - \text{daily risk free return}}{\text{daily standard deviation of returns}}$$

Many long/short strategies are self financing and the cost of capital is therefore ignored.  And instead of computing the **sharpe ratio** using returns, it is common to use the actual profits (aka: pnl).

$$\frac{\text{average daily pnl}}{\text{daily standard deviation of pnl}}$$

### Part A (1 Point)

Using the `avg` and `sdev` operator write a `sharpe` function that computes the **sharpe ratio** of a given series of daily pnl results.

<!-- grade_id: sharpe, points: 0, grade: false, solution: true -->
```q
/ compute the sharpe ratio of a list of values
sharpe:{
    / YOUR CODE HERE
    x:avg x%sdev x;
    x}

```

Given a series of daily profits (and losses) `20 -10 30 10 15 -5 20 50 30 20`, your function should return `1.037311`.  Check that it does:

```q
sharpe 20 -10 30 10 15 -5 20 50 30 20
```

Output:
```
1.037311
```

<!-- grade_id: correct_sharpe, points: 1, grade: true, solution: false -->
```q
/ check that sharpe produces the correct results
assert[0.5316765e] "e"$sharpe -4 -1 -2 -3 5 4 3 10 5 10
assert[0.4685095e] "e"$sharpe -100 100 -300 50 400 -50 1000 50 500
```

## Drawdown

Another common performance statistic for investors is the **drawdown**.  It is defined as the loss from peak to trough.  If a stock starts at 100, rises to 120 and falls to 90, the drawdown would be 30 (120-90) or 25% (30/120).  If the price rises again to 110, the drawdown would be 8.33% (10/120).  And finally, if the price moved beyond 120 to 130, the drawdown would be 0.

### Part A (1 Point)

Using the `maxs`, `|` and `%` operators, write a function `drawdown` that computes a time-series of the drawdown statistic.

<!-- grade_id: drawdown, points: 0, grade: false, solution: true -->
```q
/ compute the drawdown (as a percentage) of a list of values
drawdown:{
    / YOUR CODE HERE
    x:(h-x)%h:maxs x;
    x}

```

Given a series of prices `100 101 100 99 98 100 102 104 103`, your function should return `0 0 0.00990099 0.01980198 0.02970297 0.00990099 0 0 0.009615385`.  Check that it does:

```q
drawdown 100 101 100 99 98 100 102 104 103
```

Output:
```
0 0 0.00990099 0.01980198 0.02970297 0.00990099 0 0 0.009615385
```

<!-- grade_id: correct_drawdowns, points: 1, grade: true, solution: false -->
```q
/ check that drawdown produces the correct results
assert[0 0 0.5 0.25 0 0.15] drawdown 100 120 60 90 120 102
assert[0 0.1 0.2 0.3 0.4 0.5 0.1 0.2 0] drawdown 100 90 80 70 60 50 90 80 100
```

### Part B (1 Point)

Using the `drawdown` function, write a function `mdd` that computes the maximum drawdown of a price series.

<!-- grade_id: mdd, points: 0, grade: false, solution: true -->
```q
/ compute the maximum drawdown of a list of values
mdd:{
    / YOUR CODE HERE
    x:max drawdown x;
    x}

```

Given the same series of prices `100 101 100 99 98 100 102 104 103`, your function should return `0.02970297`.  Check that it does:

```q
mdd 100 101 100 99 98 100 102 104 103
```

Output:
```
0.02970297
```

<!-- grade_id: correct_mdd, points: 1, grade: true, solution: false -->
```q
/ check that mdd produces the correct value for a few inputs
assert[.5] mdd 100 120 110 140 100 80 70 80 90
assert[0f] mdd 100 101 102 103 104 105 106
assert[.8] mdd 100 80 60 40 30 20 80 120 60
```

## Dictionaries, Tables and Keyed Tables

Kdb+ is a vector programming language, and as such, vectors are everywhere.  "Atomic" functions work equally well on atoms and vectors.  Remember:
- A dictionary is a pair of vectors.
- A table can be treated as a list of dictionaries (though it is stored in memory as a flipped dictionary of lists). 
- A Keyed table is a pair tables and can be thought of as a dictionary mapping one table to another.

The following questions exercise your knowledge of dictionaries and tables.

### Part A (1 Point)

Using the `!` operator, create a dictionary mapping each lowercase letter of the alphabet to the numbers 1 to 26 and store it in the variable `d`.

<!-- grade_id: dictionary, points: 0, grade: false, solution: true -->
```q
d:()!()
/ YOUR CODE HERE
d:"abcdefghijklmnopqrstuvwxyz"!1+til 26;
```

```q
d
```

Output:
```
a| 1
b| 2
c| 3
d| 4
e| 5
f| 6
g| 7
h| 8
i| 9
j| 10
k| 11
l| 12
m| 13
n| 14
o| 15
p| 16
q| 17
r| 18
s| 19
t| 20
u| 21
v| 22
..
```

<!-- grade_id: correct_dictionary, points: 1, grade: true, solution: false -->
```q
/ check there are 26 elements in the dictionary
assert[26] count d
/ check the key is a string vector
assert[10h] type key d
/ check the values sum to 251
assert[351] sum d
/ check that carnegie mellon university sums to 295
assert[295] sum d "carnegie mellon university"
```

### Part B (1 Point)

Using `([]a;b;c)` notation, generate a table `t` with three columns: `date`, `sym` and `price`.  The table should have dates ranging from 2018.01.01 thru 2018.12.31, the `sym` column should all have the value `BAC`, and the `price` column should a list of random floating point numbers.

<!-- grade_id: table, points: 0, grade: false, solution: true -->
```q
t:([]a:();b:();c:())
/ YOUR CODE HERE
t:([]date:(2018.01.01+til 365);sym:(365#enlist `BAC);price:(rand each 365#100.0));
```

```q
t
```

Output:
```
date       sym price   
-----------------------
2018.01.01 BAC 39.27524
2018.01.02 BAC 51.70911
2018.01.03 BAC 51.59796
2018.01.04 BAC 40.66642
2018.01.05 BAC 17.80839
2018.01.06 BAC 30.17723
2018.01.07 BAC 78.5033 
2018.01.08 BAC 53.47096
2018.01.09 BAC 71.11716
2018.01.10 BAC 41.1597 
2018.01.11 BAC 49.31835
2018.01.12 BAC 57.85203
2018.01.13 BAC 8.388858
2018.01.14 BAC 19.59907
2018.01.15 BAC 37.5638 
2018.01.16 BAC 61.37452
2018.01.17 BAC 52.94808
2018.01.18 BAC 69.16099
2018.01.19 BAC 22.96615
2018.01.20 BAC 69.19531
..
```

<!-- grade_id: correct_table, points: 1, grade: true, solution: false -->
```q
/ check there are 365 rows in the table
assert[365] count t
/ check the three columns are date,sym,price
assert[`date`sym`price] cols t
/ check the dates range from 2018.01.01 thru 2018.12.31
assert[2018.01.01 2018.12.31] (first;last)@\:t`date
/ check the sym column has a single value: `BAC
assert[1#`BAC] distinct t`sym
/ check that the price column has floats
assert[9h] type t`price
/ check that there are no null values in the price column
assert[1b] all not null t`price
```

### Part C (1 Point)

Using `([a]b;c)` notation, generate a table `t` with three columns: andrewid, fname, lname.  The `andrewid` column should be of type symbol and the `fname` and `lname` column should be of type string (character vector). The table should have `andrewid` as its primary key. You may put as many rows as you want, but it should have the following 2 records:

```
{andrewid: jpsaris; fname: "nick"; lname: "psaris"}
{andrewid: dseppi; fname: "duane"; lname: "seppi"}
```


<!-- grade_id: keytable, points: 0, grade: false, solution: true -->
```q
t:([]a:();b:();c:())
/ YOUR CODE HERE
t:([andrewid:(`jpsaris`dseppi)];fname:("nick";"duane");lname:("psaris";"seppi"));
`t insert (`a;"sad";"sa")
t
```

Output:
```
,2

andrewid| fname   lname   
--------| ----------------
jpsaris | "nick"  "psaris"
dseppi  | "duane" "seppi" 
a       | "sad"   "sa"
```

<!-- grade_id: correct_keytable, points: 1, grade: true, solution: false -->
```q
/ check the table column names are andrewid, fname and lname
assert[`andrewid`fname`lname] cols t
/ check the primary key andrewid
assert[1#`andrewid] keys t
/ check the column types are correct
assert["sCC"] value[meta t]`t
/ check the name of jpsaris is nick psaris
assert[("nick";"psaris")] t[`jpsaris;`fname`lname]
/ check the name of dseppi is duane sepppi
assert[("duane";"seppi")] t[`dseppi;`fname`lname]
```